# Aula 3 - Exercícios Extra - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas e numpy
import pandas as pd
import numpy as np
```

### 3.9 Escreva uma função para criar uma nova coluna chamada 'segundo' que tenha o segundo maior valor de cada linha de df.


<!-- grade_id: cell-553ee248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função apply do Dataframe aplicando uma função lambda.
#            e use a função sorted dentro da função lambda. Perceba que você precisa converter o argumento da
#            da função lambda para list antes de aplicar o método sorted.
# Dica 2: Aplique a função lambda no eixo 1.

def segundo_maior(df):
    """retorne o dataframe df com uma nova coluna 'segundo' que tenha o segundo maior valor de cada linha de df
    """
    df['segundo'] = df.apply(lambda l: sorted(l, reverse=True)[1], axis=1)
    return df

```

```python
## Faça seus testes aqui
data = {'d1': [20, 36, 25,  4,  1, 20, 17, 47,  2, 40],
        'd2': [33, 25, 48, 39, 49, 37, 45, 36, 37, 31],
        'd3': [19, 47,  9, 27, 48, 26, 24, 21,  1,  2]}
df = pd.DataFrame(data)

#segundo_maior(df)
df.apply(lambda l: sorted(l, reverse=True)[1], axis=1)
```

Output:
```
0    20
1    36
2    25
3    27
4    48
5    26
6    24
7    36
8     2
9    31
dtype: int64
```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
data = {'d1': [20, 36, 25,  4,  1, 20, 17, 47,  2, 40],
        'd2': [33, 25, 48, 39, 49, 37, 45, 36, 37, 31],
        'd3': [19, 47,  9, 27, 48, 26, 24, 21,  1,  2]}
df = pd.DataFrame(data)
assert list(segundo_maior(df)['segundo']) == [20, 36, 25, 27, 48, 26, 24, 36, 2, 31]
```