# Aula 4 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
import sqlalchemy as sqla
```

### 4.8 Considere algumas linhas de uma base de dados de morbidade apresentadas a seguir. Perceba que há dados de vários meses de 2020 para cada Estado da Federação.
#### Escreva uma função para pivotar esta base de dados apresentando os meses do ano nas colunas, os municípios nas linhas e a quantidade de óbitos como valores das células. 

|       | Município               |   cod_ibge6 | mes   |   ano | mes_ano    | Óbitos   |
|------:|:------------------------|------------:|:------|------:|:-----------|:---------|
|     0 | Alta Floresta D'Oeste   |      110001 | Abr   |  2020 | 2020-04-01 | 8        |
|     1 | Alto Alegre dos Parecis |      110037 | Abr   |  2020 | 2020-04-01 | 1        |
|     2 | Alto Paraíso            |      110040 | Abr   |  2020 | 2020-04-01 | 2        |
| 33347 | Vila Boa                |      522220 | Jan   |  2020 | 2020-01-01 | -        |
| 33348 | Vila Propício           |      522230 | Jan   |  2020 | 2020-01-01 | -        |
| 33349 | Brasília                |      530010 | Jan   |  2020 | 2020-01-01 | 578      |

#### O resultado esperado é semelhante a este:

|   cod_ibge6 | 2020-01-01   |   2020-02-01 |   2020-03-01 | 2020-04-01   | 2020-05-01   | 2020-06-01   |
|------------:|-------------:|-------------:|-------------:|-------------:|-------------:|-------------:|
|      110001 |  5 |            2 |            3 |  8 | 2            | 2            |
|      110002 | 12 |           13 |           16 | 10 | 8            | 7            |
|      110003 |  - |            1 |            1 |  - | -            | 2            |
|      110004 | 23 |            8 |           29 | 11 | 17           | 13           |
|      110005 |  1 |            2 |            2 |  3 | 1            | -            |

<!-- grade_id: cell-553e52r48eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Sempre dê preferência para fazer agregação usando códigos como identificador, 
#         use nomes em ultimo caso quando não houver códigos identificadores. 
# Dica 2: Se você usar o atributo 'mes_ano' para pivotar, o pandas já ordena de forma cronológica as colunas.
def pivot_morbidade(file_path):
    """retorne um dataframe pivotado
    """
    df = pd.read_csv(path_morbidade, sep=";")
    return df.pivot(index='cod_ibge6', columns='mes_ano', values='Óbitos')

```

```python
## Faça seus testes aqui
# path_morbidade = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade_simplificado.csv'

df = pd.read_csv(path_morbidade, sep=";", na_values='-')
pivoted = df.pivot(index='cod_ibge6', columns='mes_ano', values='Óbitos')
pivoted.loc[110001].to_dict()
```

Output:
```
{'2020-01-01': 5.0,
 '2020-02-01': 2.0,
 '2020-03-01': 3.0,
 '2020-04-01': 8.0,
 '2020-05-01': 2.0,
 '2020-06-01': 2.0}
```

<!-- grade_id: cell-0e2b6tcefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_morbidade = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
test_data = {'2020-01-01': '5', '2020-02-01': '2', '2020-03-01': '3', '2020-04-01': '8', '2020-05-01': '2', 
             '2020-06-01': '2'}
assert pivot_morbidade(path_morbidade).loc[110001].to_dict() == test_data

```