# Aula 4 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
#Fazer o download dos datasets
# Se já tiver feito o download, não precisa fazer novamente.
# Se aparecer um erro dizendo que a pasta não está vazia é porque o download já foi feito
#!git clone https://github.com/alexlopespereira/curso_ciencia_dados2020.git ../../curso_ciencia_dados2020
```

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
```

### 4.1 Carregue o arquivo do IDEB 2017 no formato de um DataFrame.  Crie uma função para calcular a média do IDEB de 2017 de todos os municípios de um determinado Estado. 
#### Remova os registros cujo valor na coluna REDE sejam Municipal, Estadual ou Federal. Deixe somente os registros cujos valores na coluna REDE sejam Pública.


<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica 1: Antes de calcular, certifique-se de interpretar os valores '-' como NA durante 
#         o carregamento dos dados ou os substitua por np.nan (NA) com o metodo replace.
# Dica 2: Use a função loc para fazer o filtro de registros.
# Dica 3: Dentre várias maneiras possíveis, uma maneira bastante concisa é usar o operador & (and) para
# congregar dois critérios de filtro (seleção) de registros.
def media_ideb(df, sigla_estado):
    """retorne o valor da media do IDEB de 2017 contido na coluna col do Estado especificado por sigla_estado.
    """
    return df.loc[(df.SG_UF==sigla_estado) & (df.REDE=='Pública'), 'IDEB12_17'].mean()

```

```python
## Faça seus testes aqui
path_ideb='../../curso_ciencia_dados2020/data/originais/ideb/ideb_municipios2017.xlsx'

df_ideb = pd.read_excel(path_ideb, skiprows = 9, skipfooter=3, na_values=['-'])
```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
###################################
### ATENÇÃO PARA ESTA INSTRUÇÃO ###
# Crie na célula anterior um dataframe df_ideb. 
# Use os argumentos skiprows, na_values e skipfooter para carregar a tabela corretamente.
###################################
assert round(media_ideb(df_ideb, 'DF'),1) == 3.4
assert round(media_ideb(df_ideb, 'SP'),1) == 4.0
assert round(media_ideb(df_ideb, 'GO'),1) == 4.2
```