# Aula 2 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 2.6 Escreva uma função Python para selecionar dentre os elementos de um array aquele que mais se aproxima de um determinado valor N passado como argumento da função.
#### Se houver dois elementos equidistantes de N e forem os mais próximos de N, retorne o menor. Ou seja, para [1, 3, 4] e N=2, retorne 1.
#### Não use laços (loops), use funções Numpy.
#### Exemplo:  para um array([0, 12, 45, 34, 99]) e N=20
#### Resultado esperado: 12

```python
import numpy as np
```

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: Use as funções np.abs, ndarray.argmin e ndarray.sort 
def get_closest(arr, N):
    """retorne o elemento do array que mais se aproxima de N.
       No caso de dois elementos equidistantes, retorne o menor.
    """
    return sorted([(x, abs(x-N)) for x in arr], key=lambda x:(x[1], x[0]))[0][0]
```

```python
## Faça seus testes aqui
N=3
arr=np.array([3, 6, 1, 4])
get_closest(arr, N)
```

Output:
```
3
```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação para nota
inarr = np.array([3, 6, 1, 4])
assert get_closest(inarr, 2) == 1
```