# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.10 Escreva uma função python para substituir vírgula por ponto numa string e depois converta os números para float. Retorne a lista de números.
#### Exemplo:  "0,20 0,40 0,56 0,99"
#### Resultado esperado: [0.2, 0.4, 0.56, 0.99]

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
def convert_to_float(texto):
    """retorne uma lista de numeros do tipo float
    """
    return [float(x) for x in texto.replace(',', '.').split()]
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação para nota
text = "0,20 0,40 0,56 0,99"
assert convert_to_float(text) == [0.2, 0.4, 0.56, 0.99]
```