# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 5.3 Escreva uma função para retornar um dataframe com os nomes das colunas, o total e a porcentagem de valores NA (de cada coluna) num dataframe passado como argumento para esta função.
#### Arredonde (com o método round) o resultado para duas casas decimais (o assert vai testar com 2 casas decimais).

<!-- grade_id: cell-553e5248ebs14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função concat para concatenar duas séries. Especifique axis=1 para concatenar
#       por colunas. Se quiser dar nome as colunas depois de concatenar, utilize o argumento
#       keys=['Total','Porcentagem']
# Dica 2: Uma das séries deve conter o total de valores NA, e a outra série a 
# porcentagem de valores NA. Depois disso concatena-as usando a função concat.

def missing_percentage(df):
    """Retorna um dataframe com os nomes das colunas, o total e a porcentagem de valores
       NA em cada coluna do dataframe passado como argumento.
    """    
    total = df.isnull().sum()
    percent = round(df.isnull().sum()/len(df)*100,2)
    return pd.concat([total, percent], axis=1, keys=['Total','Porcentagem'])

```

```python
# Solução 2
def missing_percentage(df):
    tot = df.isna().sum()    
    return pd.DataFrame({
        'Total': tot,
        'Porcentagem': tot.apply(lambda t: round((t/df.shape[0])*100, 2))
    })
```

```python
## Faça seus testes aqui
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'
#path_titanic = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/titanic/titanic.xls'
df_titanic = pd.read_excel(path_titanic)

missing_percentage(df_titanic).to_dict()
```

Output:
```
{'Total': {'pclass': 0,
  'survived': 0,
  'name': 0,
  'sex': 0,
  'age': 263,
  'sibsp': 0,
  'parch': 0,
  'ticket': 0,
  'fare': 1,
  'cabin': 1014,
  'embarked': 2,
  'boat': 823,
  'body': 1188,
  'home.dest': 564},
 'Porcentagem': {'pclass': 0.0,
  'survived': 0.0,
  'name': 0.0,
  'sex': 0.0,
  'age': 20.09,
  'sibsp': 0.0,
  'parch': 0.0,
  'ticket': 0.0,
  'fare': 0.08,
  'cabin': 77.46,
  'embarked': 0.15,
  'boat': 62.87,
  'body': 90.76,
  'home.dest': 43.09}}
```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_titanic = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/titanic/titanic.xls'
test_data = {'Total': {'body': 1188,  'cabin': 1014,  'boat': 823,  'home.dest': 564,
                       'age': 263,  'embarked': 2,  'fare': 1,  'ticket': 0,  'parch': 0,  
                       'sibsp': 0,  'sex': 0,  'name': 0,  'survived': 0,  'pclass': 0},
             'Porcentagem': {'body': 90.76,  'cabin': 77.46,  'boat': 62.87,
                            'home.dest': 43.09,  'age': 20.09,  'embarked': 0.15, 
                            'fare': 0.08,  'ticket': 0.0,  'parch': 0.0,  'sibsp': 0.0,
                            'sex': 0.0,  'name': 0.0,  'survived': 0.0,  'pclass': 0.0}}

df_titanic = pd.read_excel(path_titanic)
assert missing_percentage(df_titanic).to_dict() == test_data
```