# Aula 1 - Exercícios Extra - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.13 Escreva uma função que receba como entrada uma string conforme apresentada a seguir e retorne uma lista contendo a população dos 3 Estados brasileiros mais populosos.

"""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

#### Nesta string, a informação de cada Estado está separada por um ;
 

<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função split quebra uma string numa lista baseado no caracter 
#         de separação especificado. 
#         Exemplo: '1,2,3'.split(',') retorna a lista ['1','2','3']
# Dica 2: a função .sort de uma lista ordena a referida lista


def top3_estados(lista_texto):
    """retorne os 3 estados de maior população. O parâmetro de entrada deve ser uma string
         conforme apresentado na descrição do exercício.
    """
    return sorted([int(x) for s in lista_texto.split(';') for x in s.split('=')[1:]], reverse=True)[:3]
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
lista_estados="""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;\
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;\
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;\
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;\
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;\
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;\
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

assert top3_estados(lista_estados) == [45919049, 21168791, 17264943]
```