# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.3 Escreva uma função contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    """retorne o maior item da lista
    """
    return sum([1 if len(s) > 2 and s[0]==s[-1] else 0 for s in lista]) 
        
```

```python
## Faça seus testes aqui
contagem(['abc', 'xyz', 'aba', '1221'])
```

Output:
```
2
```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
```