# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.6 Escreva uma função Python para combinar os itens de mesma chave de uma lista de  dicionários, somando os valores correspondentes.
##### Exemplo:  para [{'item': 'carro', 'quantidade': 400}, {'item': 'moto', 'quantidade': 300}, {'item': 'carro', 'quantidade': 750}]
##### Resultado esperado: {'carro': 1150, 'moto': 300}

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
def combine_items(item_list):
    """retorne um dicionário combinando (somando) as quantidades do mesmo item
    """
    ret = {}
    for item in item_list:
        ret[item["item"]] = ret.get(item["item"], 0) + item["quantidade"]
    return ret
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Validação para nota
input_data = [{'item': 'carro', 'quantidade': 400}, 
              {'item': 'moto', 'quantidade': 300}, 
              {'item': 'carro', 'quantidade': 750}]
output_result = {'carro': 1150, 'moto': 300}
assert combine_items(input_data) == output_result

```