# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 4.5 Escreva uma função para fazer o merge do dataset de morbidade consolidado com o dataset de população
#### O dataset da população deve ficar a esquerda e o dataset da morbidade a direita no resultado.

<!-- grade_id: cell-553a5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: No dataset de morbidade existem valores NA representados por caracteres específicos em várias colunas 
#         do dataframe. Descubra quais são esses caracteres e especifiquei no argumento
#         na_values. Passe nesse argumento uma lista como os caracteres que devem ser 
#         interpretados como NA.
# Dica 2: No dataset de morbidade, especifique o separador, o caracter de números decimais e o tipo do dado do 
#         cod_ibge6 como str. Ou seja, dtype={'cod_ibge6': str}.
# Dica 3: No dataset de população, converta o cod_munic e cod_uf para str no momento de carregar o dataset de população,
#         usando novamente o argumento dtype. Exatamente como foi feito no exercício da aula 3.
def merge_morbidade_pop(path_pop, path_morbidade):
    """retorne o merge dos datasets de população e morbidade.
    """
    ### BEGIN SOLUTION
    df_pop = pd.read_excel(path_pop, sheet_name="Municipios", dtype={'cod_munic': str, 'cod_uf': str})
    # Concatena o codigo da UF com o codigo do municipio
    df_pop['cod_ibge6'] = df_pop['cod_uf'] + df_pop['cod_munic'].apply(lambda x: x[:-1])
    df_morbidade = pd.read_csv(path_morbidade, sep=';', decimal=',', dtype={'cod_ibge6': str}, 
                               na_values=['-', '...'])
    return pd.merge(df_pop, df_morbidade, on='cod_ibge6') 
    ### END SOLUTION

```

```python
## Faça seus testes aqui
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'




```

<!-- grade_id: cell-0sb2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
test_data = {'Internações': 149, 'AIH_aprovadas': 149, 'Valor_total': 101282.26}
df_morbidade_pop = merge_morbidade_pop(path_pop, path_morbidade)
df_00015 = df_morbidade_pop[(df_morbidade_pop['mes']=='Abr') & (df_morbidade_pop['cod_munic']=='00015')]
df_00015 = df_00015.loc[:,['Internações', 'AIH_aprovadas', 'Valor_total']]
dict2 = df_00015.to_dict(orient='records')[0]
dict1 = df_morbidade_pop.iloc[0, 8:11].to_dict()
assert dict1 == test_data or dict2 == test_data
```