# Aula 1 - Exercícios Extra - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.12 Escreva uma função que receba uma lista de inteiros como parâmetro e retorne um dicionário com as chaves sendo os elementos da lista, e os valores sendo os elementos da mesma lista, mas na ordem inversa.
#### Exemplo: para a lista [1, 2, 3, 4, 5]
#### Resultado esperado: {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
# Dica: use a função zip. Ela retorna itens agrupados dois a dois.
def create_dict(lista):
    """retorne o maior item da lista
    """
    return dict(zip(lista, reversed(lista)))
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert create_dict([1, 2, 3, 4, 5]) == {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}
assert create_dict([-100,-1580,-1,-9100,-77]) == {-100: -77, -1580: -9100, -1: -1, -9100: -1580, -77: -100}
```