# Aula 2 - Exercícios - Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 2.2 Escreva uma função Python para somar os valores de cada linha de um array numpy 2D e depois retornar a maior destas somas
#### Exemplo:  para um array([[1  4 ], [10  8], [4  20]]) 
#### Resultado esperado: 24

<!-- grade_id: cell-f7b05dc12adbf2e2, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Dica: A função sum do numpy pode fazer o calculo da soma por linha ou por coluna, 
#       desde que o parâmetro axis seja especificado corretamente.
def max_sum(x):
    """retorne o valor maximo dentre os resultados da soma por linha de um array 2D 
    """
    return sorted([r.sum() for r in x])[-1]
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-af1b0435b6aacc9b, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a1 = np.array([[1, 4 ], [10, 8], [4, 20]])
a2 = np.array([[1,0,1,9,5],[3,1,2,4,5],[6,1,2,3,4],[8,9,2,1,7]])
assert max_sum(a1) == 24
assert max_sum(a2) == 27
```