# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
def maior(lista):
    """retorne o maior item da lista
    """
    return sorted(lista, reverse=True)[0]
```

```python
## Faça seus testes aqui
maior([1,0,-1,100,7])
```

Output:
```
100
```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert maior([1,0,-1,100,7]) == 100
assert maior([-100,-1580,-1,-9100,-77]) == -1
```