# Aula 3 - Exercícios - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe as bibliotecas pandas e numpy
import pandas as pd
import numpy as np
```

### 3.2 Escreva uma função para obter os itens que não são comuns em duas séries. Ou seja, estão na uniao mas não na intersecção.
#### Exemplo: Considere uma Serie A com os valores [1, 2, 3, 4, 5, 6, 7, 8], e
####          uma Serie B com os valores [2, 4, 6, 8, 10, 12]
#### Resultado Esperado: Uma Serie com os valores [1, 3, 5, 7,10,12]


```python
# Para carregar uma função definida em outro caderno, execute o seguinte comando
%run ../Aula3_Ex01/Aula3_Ex01.ipynb
```

<!-- grade_id: cell-553ee248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Pense na teoria de conjuntos. Como você faria isso no papel usando um diagrama de Venn.
# Dica 2: Use a função criada no exercicio 3.1
# Dica 3: Use a função np.union1d para calcular a união de duas Series.
#     Exemplo: o código pd.Series(np.union1d(t1, t2)) cria uma Serie resultante da união de t1 e t2

def complemento2(s1, s2):
    """retorne os valores que não são comuns em duas séries
    """
    u = pd.Series(np.union1d(s1, s2))
    i = pd.Series(np.intersect1d(s1, s2))
    #return u[~u.isin(i)]
    return complemento(u, i)
```

```python
# Solução sem uso do numpy (pandas "puro")
def complemento2(s1, s2):    
    return s1[~s1.isin(s2)].append(s2[~s2.isin(s1)])
```

```python
## Faça seus testes aqui
s1 = pd.Series([1, 2, 3, 4, 5, 6, 7, 8]) 
s2 = pd.Series([2, 4, 6, 8, 10, 12]) 

complemento2(s1, s2)
```

Output:
```
0     1
2     3
4     5
6     7
4    10
5    12
dtype: int64
```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
s1 = pd.Series([1, 2, 3, 4, 5, 6, 7, 8]) 
s2 = pd.Series([2, 4, 6, 8, 10, 12]) 
c = pd.Series([1, 3, 5, 7, 10, 12]) 

assert complemento2(s1,s2).reset_index()[0].equals(c) 
```