# Aula 3 - Exercícios Extra - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas e numpy
import pandas as pd
import numpy as np
```

### 3.8 Escreva uma função para ordenar uma lista de dicionários usando uma função Lambda. A função criada deve receber como argumento a chave do dicionário que será usada na ordenação da lista.


<!-- grade_id: cell-c705b9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: use a função sorted. Use o argumento key da função sorted.
def sort_dict_list(l, chave):
    """retorne a lista ordenada
    """
    return sorted(l, key=lambda i: i[chave])

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-9ffda88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
data = [{'nome': ['Pedro'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
        {'nome': ['Jose'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
       {'nome': ['Maria'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']}]

data_sorted = [{'nome': ['Jose'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
        {'nome': ['Maria'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
       {'nome': ['Pedro'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']}]

assert sort_dict_list(data, 'nome') == data_sorted
```