# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.8 Usando um defaultdict, escreva uma função Python para transformar uma lista de listas num dicionário. Os valores repetidos de uma lista devem ser associados a mesma chave do dicionário, como no exemplo abaixo.
#### Exemplo:  [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]
#### Resultado esperado: {"carro": [10, 20], "moto": [20], "trem": [5], "caminhao": [20] }

<!-- grade_id: cell-f7b05dc12adbf2e2, points: 0, grade: false, solution: true -->
```python

def dict_from_list(lista):
    """retorne o valor maximo dentre os resultados da soma por linha de um array 2D 
    """
    from collections import defaultdict
    d = defaultdict(list)
    
    for k, v in lista:        
        d[item[0]].append(item[1])

    return d
```

```python
## Faça seus testes aqui
lista1 = [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]

for item in lista1:
    print(item[0], item[1])
```

Output:
```
carro 10
moto 20
trem 5
carro 30
caminhao 20
```

<!-- grade_id: cell-af1b0435b6aacc9b, points: 1, grade: true, solution: false -->
```python
# Validação para nota
lista2 = [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]
assert dict_from_list(lista2) == {'carro': [10, 30], 'moto': [20], 'trem': [5], 'caminhao': [20]}
```