# Aula 2 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 2.4 Escreva uma função Python para converter para celcius um array de temperaturas em Fahrenheit. Arredonde o resultado até 2a casa decimal usando a função .round(decimals=2) do numpy
#### Exemplo:  para um array([0, 12, 45.21, 34, 99.91])
#### Resultado esperado: array([-17.77, -11.11, 7.33, 1.11, 37.72])

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Dica: T(°C) = (T(°F) - 32) × 5/9
def get_ceusius(F):
    """retorne um array contendo as temperaturas em graus Ceusius
    """
    return ((F - 32) * (5/9)).round(decimals=2)
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação para nota
tf = np.array([0, 12, 45.21, 34, 99.91])
tc = np.array([-17.78, -11.11, 7.34, 1.11, 37.73])
assert (get_ceusius(tf) == tc).all()
```