# Aula 1 - Exercícios Extra - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.14 Escreva uma função Python que aceite como entrada uma string de palavras separadas por vírgula e retorne uma string com as palavras em ordem alfabética e separadas por vírgula. 
#### Exemplo:  Para t="hoje,nota,teste,bolsa"
#### Resultado Esperado: "bolsa,hoje,nota,teste"

<!-- grade_id: cell-bb9537795a06bd7e, points: 0, grade: false, solution: true -->
```python
# Dica: use as funções .split e .join presente nas strings
def sort_words(words):
    """retorne uma string ordenada de palavras separadas por vírgula
    """
    return ",".join(sorted(words.split(',')))
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert sort_words("hoje,nota,teste,bolsa") == "bolsa,hoje,nota,teste"

```