# Aula 7 - Exercícios - Storytelling e Correlação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt # Biblioteca para plotar gráficos
import seaborn as sns # Biblioteca para plotar gráficos baseada no matplotlib
import re
```

### 7.2 Escreva uma função para retornar os números presentes ao final de uma string. 
#### Retorne uma lista contendo o número no formato string ou uma lista vazia se os ultimos caracteres não forem dígitos.


<!-- grade_id: cell-553qe5248eb14bc30, points: 0, grade: false, solution: true -->
```python
def get_final_numbers(s):
    """retorne uma lista contendo o número no formato string ou 
       uma lista vazia se os ultimos caracteres não forem dígitos.
    """
    return re.findall('\d+$', s)

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2bh6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert get_final_numbers('33a307a') == []
assert get_final_numbers('33a307') == ['307']
```