# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.4 Escreva uma função Python para gerar um dicionário que contenha pares (chave e valor) de números. A chave deve variar de 1 a n, e os valores sejam o valor da chave elevado ao quadrado. 
#### Exemplo:  Para n = 5
#### Resultado Esperado: {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}

<!-- grade_id: cell-bb9537795a06bd7e, points: 0, grade: false, solution: true -->
```python
# Dica: use a função range para criar a sequencia de numeros desejados
def create_dict(n):
    """retorne um dicionário com pares (x, x*x)
    """
    return {x: x**2 for x in range(1, n+1)}
    
```

```python
## Faça seus testes aqui

create_dict(3)
```

Output:
```
{1: 1, 2: 4, 3: 9}
```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert create_dict(5) == {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}

```